---
title: 'ćwiczenie 1'
description: ""
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

wykonaj polecenie w oknie obok

`@instructions`
wykonaj polecenie w oknie obok

`@hint`


`@pre_exercise_code`
```{r}
# zmienna liczbowa 
liczba <- 1

#zmienna tekstowa
tekst <- "abcs"

# pomnóż dwie zmienną liczba przez 2 i zapisz do zmiennej iloczyn
```

`@sample_code`
```{r}
# zmienna liczbowa 
liczba <- 1

#zmienna tekstowa
tekst <- "abcs"

# pomnóż dwie zmienną liczba przez 2 i zapisz do zmiennej iloczyn
```

`@solution`
```{r}
# zmienna liczbowa 
liczba <- 1

#zmienna tekstowa
tekst <- "abcs"

# pomnóż dwie zmienną liczba przez 2 i zapisz do zmiennej iloczyn
iloczyn <- liczba * 2
```

`@sct`
```{r}

```
